<a href="https://colab.research.google.com/github/Mrkomiljon/google-colab-/blob/main/vazifa(2_1_1)_arrays_(2).ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Faqatgina `YOUR CODE HERE` yoki "YOUR ANSWER HERE" deyilgan joyga dasturingizni yoki javobingizni yozing

```python
NAME = "Mukhammadiev Komiljon"
COLLABORATORS = "Mohirdev"
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
# Ma'lumotlar tahlili

## 2.1.1. N-O'lchovli massivlarga ishlov berish


```python
# Numpy kutubxonasini chaqirib olish
import numpy as np
```

### Masala(1): 1-o'lchamli massivlar yaratish (4 points)

**Masala sharti :** Elementlari ixtiyoriy  *n* ta  (*n* > 0, butun son) bo'lgan 1-o'lchamli massiv (array) yaratuvchi funksiyani tuzing. 

**Yordam :**  `array` yoki `random` funksiyasidan foydalaning. 


<!-- grade_id: one_dim, points: 0, grade: false, solution: true -->
```python
import numpy as np
def one_dim(n):
    """ n ta elementdan tashkil topgan 1-o'lchamli massiv   
    """
     # YOUR CODE HERE
m = int(input())
x= np.random.randint(1, 100, m)
x
#raise NotImplementedError()
```

Sizning funksiyangiz `n`elementlari soni, `ndim` massiv o'lchami, hamda numpy array ekanligi orqali tekshiriladi: funksiyangizni pastdagi cellda tekshirib ko'ring

<!-- grade_id: correct_one_dim, points: 4, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
for i in range(1, 10):
    import numpy as np
    assert len(one_dim(i)) == i  # n elementlar soni
    arr = one_dim(i) # massiv
    assert arr.ndim == 1 # massiv o'lchami
    assert type(arr) == np.ndarray # #numpy array 
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-6-5769e4be3c75> in <module>()
      2 for i in range(1, 10):
      3     import numpy as np
----> 4     assert len(one_dim(i)) == i  # n elementlar soni
      5     arr = one_dim(i) # massiv
      6     assert arr.ndim == 1 # massiv o'lchami

TypeError: object of type 'NoneType' has no len()
```

### Masala(2) 2-o'lchamli massivlar yaratish (6 points)

**Masala :** Shakli `shape(r, c)` bo'lgan 2-o'lchamli massiv (array) yaratuvchi funksiyani tuzing.  

**Yordam :** numpy kutubxonasida massivlar yaratuvchi funksiyalarining istalganidan foydlanishingiz mumkin. 


<!-- grade_id: two_dim, points: 0, grade: false, solution: true -->
```python
def two_dim(r, c):
    """ r ta qator, c ta ustundan tashkil topgan 2-o'lchamli massiv  
    """
     # YOUR CODE HERE
     raise NotImplementedError()
```

Funksiyangizni pastdagi cellda test qilib ko'ring

<!-- grade_id: correct_two_dim, points: 6, grade: true, solution: false -->
```python
arr = two_dim(3, 2) # massiv
assert arr.shape == (3, 2) # shape
assert arr.ndim == 2 # massiv o'lchami
assert type(arr) == np.ndarray # #numpy array 
```

```python

```